Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". When you see a line with `YOUR CODE HERE`, you have to remove the line `raise NotImplementedError()`. 

Note that you cannot edit any of the other cells.

Fill in your name and student number below:

```python
NAME = ""
STUDENT_NUMBER = ""
```

---

**Exercise 8.1**

Write a function called `top_kmers()`. The function should have three arguments: `fname`, `k`, `n_top`. The `fname` represents the FASTA filename, `k` the length of the k-mer and `n_top` the number of most abundant k-mers to return. The function should not print the k-mers, but return them as a list. The `k` and `n_top` arguments should be *optional*, with default values of 8 and 10, respectively.

**Note:** this problem set exercise has the following **extra** requirements:

* This function should ignore all FASTA file lines starting with `>`. 
* This function should return *exactly* `n_top` k-mers. If there are multiple k-mers that tie for the same place you are free to choose whichever way you like (alphabetical, first in the dictionary, etc.)
* The k-mers should be all *upper-case*.


<!-- grade_id: cell-f0beb59c46301c33, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e89fb9775084d07f, points: 5, grade: true, solution: false -->
```python
score = 0
fname = "/scratch/cfb/list-dict-loop/GSM288345_ES_Nanog_mm9.peaks.fasta"
kmers = top_kmers(fname)

try:
    assert "AACAATAG" in kmers, "AACAATAG should be the top k-mer (with 22 counts)"
    assert "TCTATTGT" in kmers, "AACAATAG should be in the top 10 (with 20 counts)"
    assert "CAAGGTCA" in kmers, "CAAGGTCA should be in the top 10 (with 20 counts)"
    assert "TGACCTTG" in kmers, "TGACCTTG should be in the top 10 (with 19 counts)"
    score += 2
except AssertionError as e:
    print(str(e.args[0]))
except Exception as e:
    print(e)
    
try:
    assert len(kmers) == 10, "10 kmers should be returned"
    assert len(kmers[0]) == 8, "kmer length should be 8"
    score += 1
except AssertionError as e:
    print(str(e.args[0]))
except Exception as e:
    print(e)

try:
    kmers = top_kmers(fname, 5, 20)
    assert len(kmers) == 20, "20 kmers should be returned"
    assert len(kmers[0]) == 5, "kmer length should be 8"
    score += 2
except AssertionError as e:
    print(str(e.args[0]))
except Exception as e:
    print(e)

if score == 0:
    raise ValueError

score
```

<!-- grade_id: cell-a9e08af10a3e011b, points: 4, grade: true, solution: false -->
```python

```

**Exercise 8.3**

The file `/scratch/cfb/files/genes.txt` contains gene annotation. Each line represents one gene. The file is tab-separated and contains the following columns:

```
1 - chromosome
2 - gene start
3 - gene end
4 - strand
5 - gene name
6 - number of exons
7 - a comma-separated list of exon sizes
8 - a comma-separated list of exon starts, relative to the gene start
```

Write a function called `gene2exons()` that reads the gene annotation from this file and converts it to a file with exon information. The function should accept two mandatory arguments, the `input_file` and the `output_file`. The output should be *tab-separated* and will contain 5 columns:

```
1 - chromosome
2 - exon start
3 - exon end
4 - exon name: gene name followed by a . and the exon number
5 - strand
```

For instance, for the gene SALL4 the output should look like this:

```
chr20    51782330    51784684   SALL4.1    -
chr20    51788860    51789141   SALL4.2    -
chr20    51790021    51792352   SALL4.3    -
chr20    51802278    51802520   SALL4.4    -
```

The gene SALL4 has four exons, which are numbered from 1 to 4. Exon numbers are not dependent on the strand, they are numbered from low start position to high start position.

<!-- grade_id: cell-8dc8fef0cfee5589, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5c3def87f913a7b2, points: 6, grade: true, solution: false -->
```python
# Only hidden tests here!

```