Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". When you see a line with `YOUR CODE HERE`, you have to remove the line `raise NotImplementedError()`. 

Note that you cannot edit any of the other cells.

Fill in your name and student number below:

```python
NAME = ""
STUDENT_NUMBER = ""
```

---

**Exercise 9.5: FASTA statistics**

Write a script called `nuc_fasta.py`.
This script should accept the name of a FASTA file as argument, read the FASTA sequences in that file and print the nucleotide content of all the sequences.
The script should print a header line, followed by the FASTA id (the sequence name) of every sequence followed by the nucleotide content of that sequence. The output should be tab-separated. The nucleotide content should be specified as a fraction of the sequence length, with two digits in the order A, C, G and T. When you run the script `nuc_fasta.py` on the input file `/scratch/cfb/scripts_debugging/sequences.fa` the output should *exactly* mach the following:

```
name	A	C	G	T
chr14:89352059-89352259	0.31	0.28	0.20	0.22
chr5:74264624-74264824	0.34	0.20	0.20	0.26
chr2:132500203-132500403	0.23	0.12	0.21	0.43
chr6:30630663-30630863	0.28	0.23	0.27	0.22
chr15_KI270905v1_alt:1999423-1999623	0.35	0.22	0.15	0.28
```

Remember that you can use an f-string for the numbers: `f"{0.1234:.2f}"` will print `0.12`.

```python
# Note: you don't write any code here as the code should be in a script
# NOTE: Call your script 'nuc_fasta.py' AND UPLOAD IT WITH YOUR NOTEBOOKS!
```

<!-- grade_id: cell-18694f29905c51ef, points: 2, grade: true, solution: false -->
```python
import os 
from tempfile import NamedTemporaryFile
import subprocess as sp

fname = "/scratch/cfb/scripts_debugging/sequences.fa"

score = 0
assert os.path.exists('nuc_fasta.py'), "no script named 'nuc_fasta.py'"
score += 1

tmpfile = NamedTemporaryFile().name
try:
    assert sp.check_call(f"python nuc_fasta.py {fname}> {tmpfile}", shell=True) == 0, "can't run the script nuc_fasta.py"
    score += 1
except Exception as e:
    print(e)
    
if score == 0:
    print("No points for this exercise. Did you use the correct name?")
    raise ValueError()

score
```

Output:
```
ERROR:root:Internal Python error in the inspect module.
Below is the traceback from this internal error.

ERROR:root:Internal Python error in the inspect module.
Below is the traceback from this internal error.

ERROR:root:Internal Python error in the inspect module.
Below is the traceback from this internal error.

Traceback (most recent call last):
  File "/home/simon/anaconda3/envs/cfb/lib/python3.8/site-packages/IPython/core/interactiveshell.py", line 3343, in run_code
    exec(code_obj, self.user_global_ns, self.user_ns)
  File "<ipython-input-1-a2d8f4b77e6a>", line 8, in <module>
    assert os.path.exists('nuc_fasta.py'), "no script named 'nuc_fasta.py'"
AssertionError: no script named 'nuc_fasta.py'

During handling of the above exception, another exception occurred:

Traceback (most recent call last):
  File "/home/simon/anaconda3/envs/cfb/lib/python3.8/site-packages/IPython/core/interactiveshell.py", line 2044, in showtraceback
    stb = value._render_traceback_()
AttributeErr
```

<!-- grade_id: cell-544e284e59af8846, points: 8, grade: true, solution: false -->
```python
# Other tests that check the output file

```