Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". When you see a line with `YOUR CODE HERE`, you have to remove the line `raise NotImplementedError()`. 

Note that you cannot edit any of the other cells.

Fill in your name and student number below:

```python
NAME = ""
STUDENT_NUMBER = ""
```

---

**Exercise 7.2: Open reading frames**

Write a function `longest_orfs` that will take a DNA sequence as input and returns the protein sequence encoded by the longest ORF for all three frames. The return value should be a dictionary, with the frame number as the key and the protein sequence as the value. This means that the dictionary will have three `key:value` pairs. The three keys are `1`, `2` and `3` and the corresponding values are the longest ORF for frame 1, the longest ORF for frame 2 and the longest ORF for frame 3. 

The stop codon (`*` in the amino acid table) should not be included in the translation of the ORF. For this exercise an open reading frame does not need to start with an ATG, it can start with any codon, except a stop codon.

For example, take the sequence `"actgcgtagagagctggagagattaggc"`. The translation of the full sequence for frame 1 would look like this:

```
TA*RAGEIR
```

In this example, the longest ORF is `RAGEIR`.

You can re-use code that you have written for exercise 4.8 and exercise 7.1. 

<!-- grade_id: cell-f7c9370aa374815c, points: 0, grade: false, solution: true -->
```python

# Adapt this function with your own code
def longest_orfs(seq):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e70609ccfbddec20, points: 1, grade: true, solution: false -->
```python
# Test simple sequence including stop codons
seq = "actgcgtagagagctggagagattaggc"
ref_result = {1: 'RAGEIR', 2: 'LRRELERLG', 3: 'CVESWRD'}
test_result = longest_orfs(seq)
for frame in ref_result:
    assert test_result[frame] == ref_result[frame], f"ORF of frame {frame} should be {ref_result[frame]}"
```

<!-- grade_id: cell-e7fcc7e972bdff40, points: 9, grade: true, solution: false -->
```python
# Hidden tests

```

**Exercise 7.5: Motif conversion**

Write a function `cons2pwm` that converts a consensus sequence into a positional weight matrix. The consensus sequence can be composed of symbols from the IUPAC DNA code. The function should ignore any non-IUPAC character. The positional weight matrix should be a two-dimensional list. Every element in the first list is a list containing the relative frequencies of A, C, G and T, in that specific order, that together sum up to 1.0.




<!-- grade_id: cell-dfe8b0355ed19e76, points: 0, grade: false, solution: true -->
```python
# Write your function here

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b95081cff4782f66, points: 10, grade: true, solution: false -->
```python
# Hidden tests

```

**Exercise 7.6: Analysis of a regulatory network**

For this question use the regulatory network represented in the following dictionary:

```python
hsc_network = {
    "Scl" : ["Scl","Gata1","Gata2","Hhex","Zfpm1","Fli1","Erg","Smad6","Runx1","Eto2"],
    "Pu.1" : ["Pu.1","Runx1","Gata1"],
    "Runx1" : ["Runx1","Pu.1","Erg"],
    "Smad6" : ["Runx1"],
    "Eto2" : ["Erg"],
    "Gata1" : ["Gata1","Scl","Pu.1","Gata2"],
    "Fli1" : ["Fli1","Runx1","Pu.1","Scl","Gata2","Hhex","Erg","Smad6"],
    "Erg" : ["Smad6","Runx1","Erg","Hhex","Gata2","Fli1"],
    "Zfpm1" : ["Gata2"],
    "Gata2" : ["Zfpm1","Runx1","Smad6","Eto2","Scl","Gata2","Hhex","Fli1","Erg"],
    "Hhex" : ["Gata2"]
}
```

Each key is a transcription factor and each value is a `list` of target genes.

Write a function `common_targets` that accepts three arguments: a TF-Target dictionary, and two TF names. The function should return a list with all the targets that are shared between the two TFs. The target list should be alphabetically ordered. When there are no common targets the function should return an empty `list`. With the network above, the following code:

```python
print(common_targets(hsc_network, "Scl", "Gata2"))
print(common_targets(hsc_network, "Smad6", "Eto2"))
print(common_targets(hsc_network, "Zfpm1", "Hhex"))
```

Should result in:
    
```python
['Erg', 'Eto2', 'Fli1', 'Gata2', 'Hhex', 'Runx1', 'Scl', 'Smad6', 'Zfpm1']
[]
['Gata2']
```

<!-- grade_id: cell-f0a5d1865f918734, points: 0, grade: false, solution: true -->
```python
# Write your function here

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bba682499f9c28b6, points: 2, grade: true, solution: false -->
```python
test_network = {
    "Scl" : ["Scl","Gata1","Gata2","Hhex","Zfpm1","Fli1","Erg","Smad6","Runx1","Eto2"],
    "Pu.1" : ["Pu.1","Runx1","Gata1"],
    "Runx1" : ["Runx1","Pu.1","Erg"],
    "Smad6" : ["Runx1"],
    "Eto2" : ["Erg"],
    "Gata1" : ["Gata1","Scl","Pu.1","Gata2"],
    "Fli1" : ["Fli1","Runx1","Pu.1","Scl","Gata2","Hhex","Erg","Smad6"],
    "Erg" : ["Smad6","Runx1","Erg","Hhex","Gata2","Fli1"],
    "Zfpm1" : ["Gata2"],
    "Gata2" : ["Zfpm1","Runx1","Smad6","Eto2","Scl","Gata2","Hhex","Fli1","Erg"],
    "Hhex" : ["Gata2"]
}

assert common_targets(test_network, "Scl", "Gata2") == ['Erg', 'Eto2', 'Fli1', 'Gata2', 'Hhex', 'Runx1', 'Scl', 'Smad6', 'Zfpm1']
assert common_targets(test_network, "Smad6", "Eto2") == []
assert common_targets(test_network, "Zfpm1", "Hhex") == ['Gata2']
```

<!-- grade_id: cell-a731c6d4287a607b, points: 8, grade: true, solution: false -->
```python
# Hidden tests

```

```python

```