Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". When you see a line with `YOUR CODE HERE`, you have to remove the line `raise NotImplementedError()`. 

Note that you cannot edit any of the other cells.

Fill in your name and student number below:

```python
NAME = ""
STUDENT_NUMBER = ""
```

---

**Exercise 5.1**

Write a function called `median()` that takes a list of numbers as an argument and returns the median as a `float` (see definition [here](https://en.wikipedia.org/wiki/Median)). 


<!-- grade_id: cell-f6bc997c61ab6ea8, points: 0, grade: false, solution: true -->
```python
def median(numbers):
    """
    Compute the median of a list of numbers.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e6cdc2cceee6dd56, points: 1, grade: true, solution: false -->
```python
# Make sure the function has the correct name and has a docstring
import inspect
assert ('median' in dir() and inspect.isfunction(median))
assert median.__doc__ is not None
assert len(median.__doc__) > 10
```

<!-- grade_id: cell-a81d2603d70b53b9, points: 1, grade: true, solution: false -->
```python
# Check if the function returns the right type of value
assert median([1.0, 2.0, 3.0]) != None, "function should use the return statement"
assert median([1.0, 2.0, 3.0]) != type(()), "function should return only one value"
assert type(median([1.0, 2.0, 3.0])) == type(0.0), "function should return a float"
assert type(median([1, 2, 3])) == type(0.0), "function should return a float"
```

<!-- grade_id: cell-44f8bffd75966ad3, points: 1, grade: true, solution: false -->
```python
# Check if the function returns the right value
assert median([1.0, 2.0, 3.0]) == 2.0, "median should be 2.0"
assert median([1.0, 3.0]) == 2.0, "median should be 2.0"
assert median([1, 3, 3, 6, 7, 8, 9]), "median should be 6.0"
assert median([1, 2, 3, 4, 5, 6, 8, 9]), "median should be 4.5"
```

<!-- grade_id: cell-f1e98dc72fb955d6, points: 1, grade: true, solution: false -->
```python
# Additional tests are run when you hand in the assignment

```

**Exercise 5.2**

Write a function called `sequence_type()` that takes one argument, a `string`. The function should return a `string` with one of three values: `"dna"`, `"protein"` or `"other"`. This should depend on the input. If the input only contains A, C, T, G or N it should return `"dna"`. If the input contains only valid one-letter IUPAC amino acid symbols (see below) it should return `"protein"`. In all other cases it should return `"other"`. The function should work, regardless of the input being upper-case, lower-case or a mix.

Note: theoretically, a sequencing consisting of A, C, G and T can be an amino acid sequence consisting of Alanines, Cysteines, Glycines and Threonines. However, you can ignore this possibility for this exercise.

```
Amino Acid Code:  Three letter Code:  Amino Acid:
----------------  ------------------  -----------
A.................Ala.................Alanine
B.................Asx.................Aspartic acid or Asparagine
C.................Cys.................Cysteine
D.................Asp.................Aspartic Acid
E.................Glu.................Glutamic Acid
F.................Phe.................Phenylalanine
G.................Gly.................Glycine
H.................His.................Histidine
I.................Ile.................Isoleucine
K.................Lys.................Lysine
L.................Leu.................Leucine
M.................Met.................Methionine
N.................Asn.................Asparagine
P.................Pro.................Proline
Q.................Gln.................Glutamine
R.................Arg.................Arginine
S.................Ser.................Serine
T.................Thr.................Threonine
V.................Val.................Valine
W.................Trp.................Tryptophan
X.................Xaa.................Any amino acid
Y.................Tyr.................Tyrosine
Z.................Glx.................Glutamine or Glutamic acid
```

<!-- grade_id: cell-dbd4ff0f03baaa8c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c584a53273381048, points: 1, grade: true, solution: false -->
```python
# Make sure the function has the correct name and has a docstring
import inspect
assert ('sequence_type' in dir() and inspect.isfunction(sequence_type)), "did you use the correct name?"
assert sequence_type.__doc__ is not None, "no docstring found"
assert len(sequence_type.__doc__) > 10
```

<!-- grade_id: cell-0f2147baf96430c4, points: 1, grade: true, solution: false -->
```python
# Check if the function returns the right type of value
assert sequence_type("X") is not None, "function should use the return statement"
assert sequence_type("X") != type(()), "function should return only one value"
assert type(sequence_type("X")) == type(""), "function should return a string"

```

<!-- grade_id: cell-40d24286d00920c3, points: 1, grade: true, solution: false -->
```python
# Check if the right type is returned
assert sequence_type("AYD") == "protein", "function should return 'protein'"
assert sequence_type("ACTG") == "dna", "function should return 'dna'"
assert sequence_type("HELLO") == "other", "function should return 'other'"
```

<!-- grade_id: cell-c2677920964b5bf8, points: 0, grade: true, solution: false -->
```python
# Additional tests are run when you hand in the assignment
```

```python

```