Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". When you see a line with `YOUR CODE HERE`, you have to remove the line `raise NotImplementedError()`. 

Note that you cannot edit any of the other cells.

Fill in your name and student number below:

```python
NAME = ""
STUDENT_NUMBER = ""
```

---

**Exercise 5.4**

Adapt your function `random_sequence()` and add an optional argument `gc`, which determines the total fraction of G + C in the sequence. This means that if you call `random_sequence(length=100, gc=0.4)` that 40 out of the 100 characters of the returned string should be a G or a C. The default GC% should be 50%. Note that there are multiple solutions possible. Hint: check the following functions from the random module:

* `randint()`
* `random()`
* `choice()`
* `shuffle()`

<!-- grade_id: cell-335da5463b03159e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-23f4c7434dbe893b, points: 1, grade: true, solution: false -->
```python
# Make sure the function has the correct name and has a docstring
import inspect
assert ('random_sequence' in dir() and inspect.isfunction(random_sequence)), "did you use the correct name for the function?"
assert random_sequence.__doc__ is not None, "function should contain a docstring"
assert len(random_sequence.__doc__) > 10, "docstring should actually contain some information"
```

<!-- grade_id: cell-952da9fc723e5505, points: 1, grade: true, solution: false -->
```python
# Check if the function returns the right type of value
assert random_sequence() != None, "function should use the return statement"
assert random_sequence()  != type(()), "function should return only one value"
assert type(random_sequence()) == type(""), "function should return a string"
assert type(random_sequence()) == type(""), "function should return a string"
```

<!-- grade_id: cell-5b6f20caf4d7d96d, points: 1, grade: true, solution: false -->
```python
# Other tests
assert len(random_sequence()) == 100, "length of returned sequence should be 100"

```

<!-- grade_id: cell-7531143f3c24797e, points: 1, grade: true, solution: false -->
```python
# Tests generating a random nucleotide sequence 

```

<!-- grade_id: cell-c55c0fdc193a0a4d, points: 2, grade: true, solution: false -->
```python
# Tests generating a certain GC%

```

**Exercise 5.5**

Write a function `list_fasta()` that returns a list of all FASTA files in a directory. The first argument to `list_fasta()` should be the directory. For this exercise you can assume that FASTA files have a `.fa`, a `.fsa` or a `.fasta` extension.

<!-- grade_id: cell-85fe92391dde7270, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5de8dafcaf8a525a, points: 7, grade: true, solution: false -->
```python
# A variety of tests are run

```

```python

```